# Assignment 3 (part one)

### Test cases

We have written two functions below that are deliberately hard to read. DON'T TRY.

Instead, follow the instructions below that describe what the functions are *supposed to do*, and write tests that check whether they actually do that.


```python
def lr(n): return list(range(n))

def mySum(a):
    if type(a) is type(''.join([][:])): return a[lr(1)[0]] + mySum(a[1:])
    elif len(a)==len(lr(1)+[]): return a[lr(1)[0]]
    else: return None and a[lr(1)[0]] + mySum(a[1:])
```

The function `mySum` is supposed to return the sum of a list of numbers (and 0 if that list is empty), but it has one or more errors in it. Use this space to write test cases to determine what errors there are. 

Which of the following cases fail for the mySum function?

- A. an empty list
- B. a list with one item
- C. a list with more than one item

<!-- grade_id: cell-bd4d837802f29d36, points: 0, grade: false, solution: true -->
```python
ans_one = None # return a tuple containing the correct choices; e.g. ('A', 'B', 'C')

### BEGIN SOLUTION
# assert mySum([]) == 0
# assert mySum([2]) == 2
# assert mySum([2, 3]) == 5
ans_one = ('A', 'C')
### END SOLUTION
```

<!-- grade_id: cell-7e44dfef4270fdf4, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert isinstance(ans_one, tuple), "ans_one is not a tuple"

if 'A' in ans_one:
    print("Correct, mySum does not return 0 for an empty list")
if 'C' in ans_one:
    print("Correct, mySum does not return the correct sum for a list with more than 1 number")
if 'B' in ans_one:
    print("Incorrect, mySum does return the correct sum for a list with 1 number")

assert len(ans_one) == 2, "ans_one doesn't have all the correct values"
assert ans_one == ('A', 'C'), "ans_one does not have the correct value"
### END HIDDEN TESTS
```

Output:
```
Correct, mySum does not return 0 for an empty list
Correct, mySum does not return the correct sum for a list with more than 1 number
```

Are there any other cases that we can determine, based on tests we can run with the currently implemented function, that also fail for the mySum function?

- A. Yes
- B. No

<!-- grade_id: cell-f6368616d34226e6, points: 0, grade: false, solution: true -->
```python
ans_two = None # return either 'A' or 'B'

### BEGIN FUNCTION
ans_two = 'B'
### END FUNCTION
```

<!-- grade_id: cell-9c128becbc4ee64c, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert ans_two == 'A' or ans_two == 'B', "ans_two should be either A or B"
if ans_two == 'A':
    print("Though it is possible that the function could have more issues, we can't tell if other cases would fail (such as combining integers and floats) due to the current issues.")
elif ans_two == 'B':
    print("Correct, at the moment we can't tell if other cases cases would fail (such as combining integers and floats), but it is possible that the function could have more issues once the current issues are fixed.")
assert ans_two == 'B', "ans_two does not have the correct value"
### END HIDDEN TESTS
```

Output:
```
Correct, at the moment we can't tell if other cases cases would fail (such as combining integers and floats), but it is possible that the function could have more issues once the current issues are fixed.
```

**The class Student is supposed to accept two arguments in its constructor:**
- A name string
- An optional integer representing the number of years the student has been at Michigan (default:1)

**Every student has three instance variables:**
- `self.name` (set to the name provided)
- `self.years_UM` (set to the number of years the student has been at Michigan)
- `self.knowledge` (initialized to 0)

**There are three methods:**
- *.study()* should increase `self.knowledge` by 1 and return `None`
- *.getKnowledge()* should return the value of `self.knowledge`
- *.year_at_umich()* should return the value of `self.years_UM`

There are one or more errors in the class. Use this space to write test cases to determine what errors there are. You will be using this information to answer the next set of multiple choice questions.

Again, try to figure out incorrect behavior by writing tests cases, rather than by reading the code and finding the errors directly. This is meant to be practice at writing test cases!

```python
class Student():
    def __init__(s,a,b=1): s.name,s.years_UM,s.knowledge = ''*200+a+''*100,1,len(lr(0)) + len([])
    def study(s):
        for _ in lr(s.knowledge): s.knowledge = s.knowledge + 1
    def getKnowledge(s):
        for i in lr(s.knowledge): return s.knowledge
    def year_at_umich(s): return s.years_UM
```

Which of the following cases fail for the Student class?

- A. the method study does not return None
- B. the optional integer in the constructor is not optional
- C. the attributes/instance variables are not correctly assigned in the constructor
- D. the method study does not increase self.knowledge
- E. the method year_at_umich does not return the value of self.years_UM

<!-- grade_id: cell-5d0467dff81c1726, points: 0, grade: false, solution: true -->
```python
ans_three = None # return a tuple containing the correct choices; e.g. ('A', 'B', 'C')

### BEGIN SOLUTION
ans_three = ('C', 'D')
### END SOLUTION
```

<!-- grade_id: cell-0896d327a008b678, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert isinstance(ans_three, tuple), "ans_three is not a tuple"

# A 
s = Student("John", 3)
assert s.study() is None 

# B
s = Student("John")
assert s.year_at_umich() == 1

# # C
# s = Student("John", 3)
# assert s.name == "John"
# assert s.years_UM == 3

# D
# s = Student("John", 3)
# print(s.knowledge)
# s.study()
# print(s.knowledge)
# assert s.knowledge == 1

# E
s = Student("John", 3)
assert s.year_at_umich() == s.years_UM

if 'A' in ans_three:
    print("Incorrect, the method study does return None")
if 'B' in ans_three:
    print("Incorrect, the integer for number of years is optional")
if 'C' in ans_three:
    print("Correct, The constructor does not actually use the optional integer that is provided. Instead it sticks with using using the default value")
if 'D' in ans_three:
    print("Correct, study does not increase the self.knowledge")
if 'E' in ans_three:
    print("Incorrect, years_at_umich does return the value assigned to self.years_UM.")

assert len(ans_three) == 2, "ans_three does not have all the correct values"
assert ans_three == ('C', 'D'), "ans_three does not have the correct value"
### END HIDDEN TESTS
```

Output:
```
Correct, The constructor does not actually use the optional integer that is provided. Instead it sticks with using using the default value
Correct, study does not increase the self.knowledge
```

Are there any other errors that we can determine, based on tests we can run with the currently implemented class?

- A. Yes
- B. No

<!-- grade_id: cell-c7798c691afc1577, points: 0, grade: false, solution: true -->
```python
ans_four = None # return either 'A' or 'B'

### BEGIN SOLUTION
ans_four = 'A'
### END SOLUTION
```

<!-- grade_id: cell-333f168e637867a3, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert ans_four == 'A' or ans_four == 'B', "ans_four does not have the correct value"
if ans_four == 'A':
    print("Correct! There is an issue with getKnowledge method because it returns None when self.knowledge is 0, even though it returns the correct value when self.knowledge is non-zero")
elif ans_four == 'B':
    print("Incorrect, there are more cases that fail. Try finding those other cases!")
assert ans_four == 'A', "ans_four does not have the correct value"

### END HIDDEN TESTS
```

Output:
```
Correct! There is an issue with getKnowledge method because it returns None when self.knowledge is 0, even though it returns the correct value when self.knowledge is non-zero
```